***Support Vector Machines***
    Using SVM with Python to predict whether a breast cancer tumor is malignant or benign. Exercise uses numpy, pandas, and scikitlearn and utilizes train test split, SVM, SVC,  and GridSearch to identify the best parameters for prediction.

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline
```

```python

```

<!-- grade_id: cell-bbc3dc112799b91f, points: 0, grade: false, solution: true -->
```python
# 1 point

# import and load the breast cancer dataset 
from sklearn import datasets
# output a variable where the data is loaded straight from sklearn



# YOUR CODE HERE
cancer = datasets.load_breast_cancer()

```

<!-- grade_id: cell-89604aed63202fe9, points: 1, grade: true, solution: false -->
```python
assert cancer.data.shape == (569,30)
```

```python
print(cancer['DESCR'])
```

Output:
```
.. _breast_cancer_dataset:

Breast cancer wisconsin (diagnostic) dataset
--------------------------------------------

**Data Set Characteristics:**

    :Number of Instances: 569

    :Number of Attributes: 30 numeric, predictive attributes and the class

    :Attribute Information:
        - radius (mean of distances from center to points on the perimeter)
        - texture (standard deviation of gray-scale values)
        - perimeter
        - area
        - smoothness (local variation in radius lengths)
        - compactness (perimeter^2 / area - 1.0)
        - concavity (severity of concave portions of the contour)
        - concave points (number of concave portions of the contour)
        - symmetry
        - fractal dimension ("coastline approximation" - 1)

        The mean, standard error, and "worst" or largest (mean of the three
        worst/largest values) of these features were computed for each image,
        resulting in 30 features.  For instance, field 0 is Mean Radi
```

```python
cancer['feature_names']
```

Output:
```
array(['mean radius', 'mean texture', 'mean perimeter', 'mean area',
       'mean smoothness', 'mean compactness', 'mean concavity',
       'mean concave points', 'mean symmetry', 'mean fractal dimension',
       'radius error', 'texture error', 'perimeter error', 'area error',
       'smoothness error', 'compactness error', 'concavity error',
       'concave points error', 'symmetry error',
       'fractal dimension error', 'worst radius', 'worst texture',
       'worst perimeter', 'worst area', 'worst smoothness',
       'worst compactness', 'worst concavity', 'worst concave points',
       'worst symmetry', 'worst fractal dimension'], dtype='<U23')
```

<!-- grade_id: cell-13eab9c14ca47094, points: 0, grade: false, solution: true -->
```python
# 1 point

#load the data from the dataset into a pd df

#output a dataframe with the data of the dataset and the column names as feature names
#IT SHOULD NOT INCLUDE THE TARGET 

df_feat = pd.DataFrame(cancer.data,columns=cancer.feature_names)
# YOUR CODE HERE

```

<!-- grade_id: cell-1c2998ddc5440385, points: 1, grade: true, solution: false -->
```python
b = True
for i in cancer['feature_names']:
    if df_feat[i].shape != (569,):
        b = False
        break
assert b
```

```python
cancer.target.shape
```

Output:
```
(569,)
```

<!-- grade_id: cell-58bf6e127f5d9c62, points: 0, grade: false, solution: true -->
```python
# 1 point

# create a dataframe with all the targets and the column name your be 'Cancer'
# output a single column dataframe
df_target = pd.DataFrame(cancer.target, columns=['Cancer'])

# YOUR CODE HERE

```

<!-- grade_id: cell-c504ee25e1401818, points: 1, grade: true, solution: false -->
```python
assert df_target['Cancer'].shape == (569,)
```

<!-- grade_id: cell-cea569421704556d, points: 0, grade: false, solution: true -->
```python
#1 point
# get the train test split, tested with random state 0 and test_size 0.2
from sklearn.model_selection import train_test_split
X,y = cancer.data , cancer.target
x_train,x_test,y_train,y_test = train_test_split(X,y,test_size=0.2,random_state=0)
# all the subsets should be numpy arrays
# TIP carefull with the shape of the Y
# YOUR CODE HERE

```

<!-- grade_id: cell-8557f6d20b5ebd4a, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (455,30)
assert x_test.shape == (114,30)
assert y_train.shape == (455,)
assert y_test.shape == (114,)
```

<!-- grade_id: cell-d569929ec688ee8f, points: 0, grade: false, solution: true -->
```python
# 2 points
#Train the support vector classifier
# import if needed
from sklearn.svm import SVC
clf = SVC()
clf.fit(x_train,y_train)
```

Output:
```
SVC()
```

<!-- grade_id: cell-741670633ab66a5b, points: 2, grade: true, solution: false -->
```python
assert clf.support_vectors_.shape == (123,30)
```

<!-- grade_id: cell-ee4e7bf6bb7b13e9, points: 0, grade: false, solution: true -->
```python
# 3 points
# predict the testing data and get the accuracy
from sklearn import metrics
# output preds is the array with the predictions and acc is a float between 0 and 1
preds = clf.predict(x_test)
acc = metrics.accuracy_score(y_test,preds)
print('Accuracy is:', acc)
# YOUR CODE HERE


```

Output:
```
Accuracy is: 0.9298245614035088
```

```python
from sklearn.metrics import confusion_matrix



conf_matrix = confusion_matrix(y_test, preds)
print(conf_matrix)
```

<!-- grade_id: cell-a8b6a273210f6822, points: 2, grade: true, solution: false -->
```python
assert acc > 0.85
assert preds.shape == (114,)
```

***Everything is being classified into a single class! Use Grid Search to search for the best parameters! (Scikitlearn GridSearchCV)***


<!-- grade_id: cell-43d48d4bf75c556e, points: 0, grade: false, solution: true -->
```python
# 2 points
#Set parameters
# the parameters should be in a dict format
params = {'C':np.linspace(1,10,10).astype(int), 'kernel':['linear',  'sigmoid']      }
# YOUR CODE HERE

```

<!-- grade_id: cell-28a594d468d2e212, points: 2, grade: true, solution: false -->
```python
assert len(params) == 2
assert len(params['kernel']) == 2
```

<!-- grade_id: cell-8ac1a0df91772743, points: 0, grade: false, solution: true -->
```python
#Fit grid (Note to self: first it runs the same loop with cross-validation to find the best parameter combination. Once it has the best combination, it runs fit again on all data passed to fit (without cross-validation), to built a single new model using the best parameter setting.)
# 4 points
# add an svm to the grid and train it
# the SVM should be an untrained one
from sklearn.model_selection import GridSearchCV
svc = SVC(random_state=0)
grid_svc = GridSearchCV(svc,params)
grid_svc.fit(x_train,y_train)

# YOUR CODE HERE


```

Output:
```
GridSearchCV(estimator=SVC(random_state=0),
             param_grid={'C': array([ 1,  2,  3,  4,  5,  6,  7,  8,  9, 10]),
                         'kernel': ['linear', 'sigmoid']})
```

```python
tuned_model =grid_svc.best_estimator_
print(tuned_model)
print(grid_svc.best_score_)
```

Output:
```
SVC(C=1, kernel='linear', random_state=0)
0.9604395604395604
```

```python
# pd.DataFrame(grid_svc.cv_results_)
```

Output:
```
mean_fit_time  std_fit_time  mean_score_time  std_score_time param_C  \
0        1.321381      0.531537         0.001400    4.897264e-04       1   
1        0.009194      0.000980         0.002199    4.004241e-04       1   
2        1.056346      0.509440         0.000400    4.898042e-04       2   
3        0.009194      0.000979         0.001999    1.288343e-06       2   
4        2.595992      0.645078         0.001000    4.623108e-07       3   
5        0.015391      0.001959         0.003199    4.009010e-04       3   
6        3.732489      0.982939         0.001200    4.000427e-04       4   
7        0.008994      0.000895         0.002199    4.001644e-04       4   
8        4.648123      3.045752         0.001000    6.323345e-04       5   
9        0.009394      0.000489         0.002399    7.991793e-04       5   
10       4.591957      3.141175         0.000999    1.256174e-06       6   
11       0.009193      0.000979         0.003199    1.165370e-03       6   
12       5.57474
```

<!-- grade_id: cell-085b7fafdce47092, points: 4, grade: true, solution: false -->
```python
dic = grid_svc.get_params(deep=False)
assert type(dic['estimator']) == type(SVC()) 
assert dic['pre_dispatch'] == '2*n_jobs'
```

<!-- grade_id: cell-7901a09d1a9092ac, points: 0, grade: false, solution: true -->
```python
# 2 point
#Get in a list the best parameters of the grid
#output python list with the best parameters sorted
best_params = sorted(grid_svc.cv_results_)


```

<!-- grade_id: cell-99bb42084e0ae89b, points: 2, grade: true, solution: false -->
```python
assert best_params[0] == 'mean_fit_time'
assert best_params[1] == 'mean_score_time'
assert best_params[2] == 'mean_test_score'
assert best_params[3] == 'param_C'
```

<!-- grade_id: cell-4eb30db85584892f, points: 0, grade: false, solution: true -->
```python
# 5 points
# predict the testing set and get the accuracy
# accuracy as always will be between 0 and 1

best_model = tuned_model
best_model.fit(x_train,y_train)
preds = best_model.predict(x_test)
acc = metrics.accuracy_score(y_test,preds)
print('Accuracy is:', acc)

# YOUR CODE HERE

```

Output:
```
Accuracy is: 0.956140350877193
```

<!-- grade_id: cell-67663369ff5f55b0, points: 5, grade: true, solution: false -->
```python

assert acc > 0.92
assert preds[preds.argmax()] == 1
assert preds.shape == (114,)
```

<!-- grade_id: cell-b467b5e84dfff2d3, points: 0, grade: false, solution: true -->
```python
# 3 points
# get the confusion matrix
#import if needed




conf_mat = confusion_matrix(y_test, preds)
print(conf_mat)
# YOUR CODE HERE



```

Output:
```
[[46  1]
 [ 4 63]]
```

<!-- grade_id: cell-fd7017a6eb94325a, points: 3, grade: true, solution: false -->
```python
assert conf_mat.shape == (2,2)
assert conf_mat[0][1] < 7
assert conf_mat[1][0] < 3

```

Output:
```
AssertionError                            Traceback (most recent call last)
c:\Users\T430\Desktop\Desktop\Strivestuff\stive_exercises\Chapter 02\12. SVM\ml7.ipynb Cell 32' in <cell line: 3>()
      <a href='vscode-notebook-cell:/c%3A/Users/T430/Desktop/Desktop/Strivestuff/stive_exercises/Chapter%2002/12.%20SVM/ml7.ipynb#ch0000027?line=0'>1</a> assert conf_mat.shape == (2,2)
      <a href='vscode-notebook-cell:/c%3A/Users/T430/Desktop/Desktop/Strivestuff/stive_exercises/Chapter%2002/12.%20SVM/ml7.ipynb#ch0000027?line=1'>2</a> assert conf_mat[0][1] < 7
----> <a href='vscode-notebook-cell:/c%3A/Users/T430/Desktop/Desktop/Strivestuff/stive_exercises/Chapter%2002/12.%20SVM/ml7.ipynb#ch0000027?line=2'>3</a> assert conf_mat[1][0] < 3

AssertionError:
```

# Discussion of results

### Using the default values of SVC to train the data an accuracy of about 92% was gotten 
### But performing a gridsearch to get the best parameters to use for the clf, I got the best paramter to be (C=1, kernel='linear', random_state=0)
### Using the above to train the model, there was an increase in the accuracy from 92% to about 96%