# Pandas NBgrader exercises

In this notebook you will find the 6 deliverable exercises regarding pandas. This notebook is autograded so please dont modify any cell except to add your solution. Modifying this notebook may result in you not getting a grade.

```python
import pandas as pd
import numpy as np
```

```python
data = {'animal': ['cat', 'cat', 'snake', 'dog', 'dog', 'cat', 'snake', 'cat', 'dog', 'dog'],
        'age': [2.5, 3, 0.5, np.nan, 5, 2, 4.5, np.nan, 7, 3],
        'visits': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'priority': ['yes', 'yes', 'no', 'yes', 'no', 'no', 'no', 'yes', 'no', 'no']}

labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']

df = pd.DataFrame(data, index=labels)

```

**13.** Change the age in row 'f' to 1.5.

<!-- grade_id: cell-f8d4a64dd0fb6112, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.loc['f','age'] = 1.5
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/3734779232.py in <module>
      1 # YOUR CODE HERE
      2 df.loc['f','age'] = 1.5
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-675b0274a23f22c9, points: 1, grade: true, solution: false -->
```python
assert df['age']['f'] == 1.5
```

**18.** Sort `df` first by the values in the 'age' in *decending* order, then by the value in the 'visit' column in *ascending* order.



<!-- grade_id: cell-bb44e438722da485, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df=df.sort_values(['age','visits'],ascending=(False,True))
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/3999439356.py in <module>
      1 # YOUR CODE HERE
      2 df=df.sort_values(['age','visits'],ascending=(False,True))
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-9fa22615a21cf3f1, points: 1, grade: true, solution: false -->
```python
b = True
for i in range(df.shape[0]-1):
    if df['age'].values[i] < df['age'].values[i+1]:
        b = False
        break
    
assert b
    
```

**23.** Given a DataFrame of numeric values, say
```python
df = pd.DataFrame(np.random.random(size=(5, 3))) # a 5x3 frame of float values
```

how do you subtract the row mean from each element in the row?

<!-- grade_id: cell-5137285459670e80, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = pd.DataFrame(np.random.random(size=(5, 3)))
df.sub(df.mean(axis=1), axis=0)
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/4069995844.py in <module>
      2 df = pd.DataFrame(np.random.random(size=(5, 3)))
      3 df.sub(df.mean(axis=1), axis=0)
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-8ceb24481d7d7a28, points: 2, grade: true, solution: false -->
```python
assert df.shape == (5,3)
assert df[0].max() <= 1
assert df[0].min() >= -1
```

**24.** Suppose you have DataFrame with 10 columns of real numbers, for example:

```python
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
```
Which column of numbers has the smallest sum? (Find that column's label.)

<!-- grade_id: cell-7830a1e1958310b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
df.sum().idxmin()
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/1124236102.py in <module>
      2 df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
      3 df.sum().idxmin()
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-a0daad4562804d2a, points: 2, grade: true, solution: false -->
```python
assert idx in 'abcdefghij'
```

Output:
```
NameError                                 Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/277311637.py in <module>
----> 1 assert idx in 'abcdefghij'

NameError: name 'idx' is not defined
```

The next three puzzles are slightly harder...

**26.** You have a DataFrame that consists of 10 columns of floating--point numbers. Suppose that exactly 5 entries in each row are NaN values. For each row of the DataFrame, find the *column* which contains the *third* NaN value.

(You should return a Series of column labels.)

<!-- grade_id: cell-e4f640a69590b406, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-182477a2beb50458, points: 3, grade: true, solution: false -->
```python
assert indexes == [3, 8, 5, 6, 8]
assert res.shape == (5,)
assert list(res.values) == [3, 8, 5, 6, 8]
```

**27.** A DataFrame has a column of groups 'grps' and and column of numbers 'vals'. For example: 

```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
```
For each *group*, find the sum of the three greatest values.

<!-- grade_id: cell-da30c3e27df05854, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
data1 = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
df = pd.DataFrame(data1)
df.groupby('grps')['vals'].nlargest(3).sum(level=0)
raise NotImplementedError()
```

Output:
```
C:\Users\T430\AppData\Local\Temp/ipykernel_13500/107053083.py:5: FutureWarning: Using the level keyword in DataFrame and Series aggregations is deprecated and will be removed in a future version. Use groupby instead. df.sum(level=1) should use df.groupby(level=1).sum().
  df.groupby('grps')['vals'].nlargest(3).sum(level=0)
NotImplementedError                       Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/107053083.py in <module>
      4 df = pd.DataFrame(data1)
      5 df.groupby('grps')['vals'].nlargest(3).sum(level=0)
----> 6 raise NotImplementedError()

NotImplementedError:
```

```python
assert df.values[0] == 409
assert df.values[1] == 156
assert df.values[2] == 345
```

Output:
```
ValueError                                Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/3700271275.py in <module>
----> 1 assert df.values[0] == 409
      2 assert df.values[1] == 156
      3 assert df.values[2] == 345

ValueError: The truth value of an array with more than one element is ambiguous. Use a.any() or a.all()
```

<!-- grade_id: cell-dd5c34897ebaa573, points: 3, grade: true, solution: false -->
```python
assert res.values[0] == 409
assert res.values[1] == 156
assert res.values[2] == 345
```

Output:
```
NameError                                 Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_13500/2384304782.py in <module>
----> 1 assert res.values[0] == 409
      2 assert res.values[1] == 156
      3 assert res.values[2] == 345

NameError: name 'res' is not defined
```

```python

```